```python
# Initialize Otter
import otter
grader = otter.Notebook("HW4.ipynb")
```

# Homework 4

This notebook contains Python problems you need to complete. Write your codes directly here. Write your codes directly here. Save your code and upload this file ("HW3.ipynb") to [Gradescope](https://www.gradescope.com). 

Note that there are some hidden tests, so you may not get full points if these tests fail. Do not be misled by the message "q# passed!".

### This assignment must be completed individually.
- Review the syllabus to avoid any academic integrity issues. 
- All codes in this Jupyter notebook file should be your own creation. 
- You are encouraged to openly collaborate on **Piazza** with your peers. Therefore, any code snippets shared on Piazza will not be subject to plagiarism even if code similarity checking tool discovers identical codes or similar coding patterns.
- You may lose 2 points for each error that prevents Gradescope from auto-grading your submission.  It would be wise to comment out the lines that cause errors if you cannot fix them.

### Common errors and how to fix/prevent them
The following errors seem to be common among students. So, tips on how to fix or prevent them are given here.
- typos:  Take advantage of auto completion using the \[TAB\] key.
- incorrect use of symbols: Be careful match the opening and closing quotations, parentheses, brackets, braces. Using too many parentheses often create syntax errors.  Create some temporary variables may reduce such errors.
- wrong indentation: the error message may provide where it occured.
- incorrect data type: Use the type conversion functions, such as int(), float(), str(). I have seem many people not wrapping their name and id with quotation marks for the pledge part.
- other syntax errors: Read the error message and take a look at where it occured.  Most of the time, error messages help you fix the error.
- logic errors: Even when no syntax error is found, your answer can be marked wrong if you have logic errors. Sometimes, you can get hints from grader.check() output. Wrong indentation oftentimes leads to a logic error. The best way to fix would be inserting testing codes such as printing object, calling the function you create with sample inputs. Fixing logic errors is a challenge for even professional coders.
- submitting a wrong file to Gradescope: Make sure you follow the instructions carefully. Know how to browse to a file.  Also, try not to create multiple files and check the time stamp if you do.

### Other tips
- Restart the Kernel (select Kernel>Restart menu) if your notebook hangs and do not run your code.  After restarting the kernel, you may want to run all cells above your current cell (Cell > Run All Above) menu, to make sure all variables/objects you need are created.
- Once you are ready to submit your file, selecting **Kernel > Restart All** menu will let you rerun all codes from top to bottom.

```python
%pprint # to turn off pretty printing
```

### Data Manipulation with Pandas

For the questions below, you also need the data files "LoanDataOriginal.csv" and "covid_all.csv".  So, make sure they are in the *same* folder where this file is.

```python
# This will tell you what your working directory is
%pwd
```

```python
# If you do not see LoanDataOriginal.csv and covid_all.csv, you need to use %cd '***The correct directory path***'
%ls
```

```python
# import numpy and pandas
import numpy as np
import pandas as pd
```

```python
# Import the file as DataFrame object and assign it to loans. 
# The index column will be generated automatrically.
loans=pd.read_csv("LoanDataOriginal.csv")
loans.head()
```

```python
# Check the structure of loans DataFrame object.
loans.info()
```

```python
# Descriptive statistics of loans
loans.describe()
```

**Question 1 (1pt):** 

After reading "LoanDataOriginal.csv" as a pandas.core.frame.DataFrame object named loan (already coded above), change the data type of `LoanStatus` as `category`. Refer to the pandas documentation for the [category](https://pandas.pydata.org/pandas-docs/stable/user_guide/categorical.html) data type and check out the examples. 

```python
# Write your code here.
loans.LoanStatus = loans.LoanStatus.astype('category')
```

```python
# See if LoanStatus has the correct dtype
print(loans['LoanStatus'].dtype)
loans['LoanStatus'].dtype
```

```python
grader.check("q1")
```

**Question 2 (4pts):** 
- Replace all values of `99999999` in `CurrentLoanAmount` with `NaN`.  That is, assign `np.NaN` to all items in column `CurrentLoanAmount` that satisfies the condition `loans['CurrentLoanAmount']==99999999`. 
- Rename the column name of `CurrentLoanAmount` to `LoanAmount` and that of `CurrentCredit` to `Credit`.
- Take the average of `LoanAmount` and store it in a variable named `avg_loan`.  The `np.mean()` function or Series object's `mean()` method will ignore the NaNs.
- Replace all NaNs in `LoanAmount` with `avg_loan`. 

```python
# Write your code here.
loans.loc[loans.CurrentLoanAmount==99999999,'CurrentLoanAmount'] = np.NaN #rename columns CurrentLoanAmount and CurrentCredit to 'LoanAmount' and 'Credit' respectively
loans = loans.rename(columns={'CurrentLoanAmount': 'LoanAmount', 'CurrentCredit': 'Credit'}) #get average of LoanAmount
avg_loan = loans.LoanAmount.mean() #replace Nan values in LoanAmount with 'avg_loan'
loans['LoanAmount'] = loans['LoanAmount'].fillna(avg_loan)
```

```python
grader.check("q2")
```

**Question 3 (3pts):** 

- Add a new column `Amt_per_Acct` to the DataFrame object `loans`. `Amt_per_Acct` should be computed as `LoanAmount / OpenAccounts`.  
- Create a DataFrame (subset) that contains the loans that are fully paid and assign it to a variable named `fully_paid`. Include all columns.
- (Optional for 1pt extra credit) Create a data frame named open_le_10 by taking the first 4 columns and the last 4 columns for loans of customers with 10 or less OpenAccounts.

```python
# Write your code here.
loans['Amt_per_Acct'] = loans['LoanAmount'] / loans['OpenAccounts'] #create a subset called 'fully_paid' with all columns and LoanStatus=Fully Paid
fully_paid = loans[loans.LoanStatus=='Fully Paid']
```

```python
# Check if the new column has been added.
loans.head(3)
```

```python
grader.check("q3")
```

**Question 4 (2pts):** 

From the `loans` DataFrame object: 
- Create a DataFrame object named `high_risk` which has selected columns `[CustomerID, LoanAmount`, `Credit`, `Income`,`CreditScore]` and includes loans of customers with `YearsCreditHistory` less than average and `MonthlyDebt` greater than average.  Remember that you need to use bitwise logical operators (`&`,`|`, `~`) instead of regular logical operators (`and`, `or`, `not) when comparing the elements of an array-like structure.
- Create a data frame named valued_customer with selected columns - `LoanID, LoanStatus, OpenAccounts, Amt_per_Acct, MonthlyDebt` - for fully paid loans where Amt_per_Acct is \$125,000 or more.

```python
# Write your code here.
avg_creditHistory = loans['YearsCreditHistory'].mean()
avg_monthlydebt = loans['MonthlyDebt'].mean()
list_col = ['CustomerID','LoanAmount','Credit','Income','CreditScore']

high_risk = loans[(loans['YearsCreditHistory']<avg_creditHistory) & (loans['MonthlyDebt']>avg_monthlydebt)][list_col]

list_col = ['LoanID','LoanStatus','OpenAccounts','Amt_per_Acct','MonthlyDebt']

valued_customer = fully_paid[fully_paid.Amt_per_Acct>=125000][list_col]
```

```python
# Display and check if your answer works.
print(high_risk.head(3))
valued_customer.head(3)
```

```python
grader.check("q4")
```

**Question 5 (3pt):** 

- Import "covid_all.csv" as a pandas.core.frame.DataFrame object named `covid_all`. Use the first column as the index column. 

- Remove all rows that are not the data for the 50 states and DC and store the remaining rows in a new variable named `covid_main`. Hint: Use columns `state` and `Type` to include only the states and DC (i.e., `covid_all['state']=='DC' | covid_all['Type']=='State'`). You just need to use `loc` to apply this filter and select all columns (using :).

- Select the columns, `date`, `state`, `Name`, `Population`, `deathIncrease`, `hospitalizedCurrently`, `positiveIncrease`, and `totalTestResultsIncrease` from `covid_main`. Then, rename these columns with new names `Date`, `StateCode`, `StateName`, `Pop`, `NewDeath`, `HospNow`, `NewCase`, and `NewTest`. Finally, assign this to a new variable named `covid`.

- Convert the `Date` column to datetime. The code for this is provided for you below.

```python
# Write your code here.
covid_all = pd.read_csv("covid_all.csv", index_col=0)

covid_main = covid_all.loc[(covid_all['state']=='DC') | (covid_all['Type']=='State'),:]

covid=covid_main.loc[:,['date', 'state', 'Name', 'Population', 'deathIncrease',
'hospitalizedCurrently', 'positiveIncrease', 'totalTestResultsIncrease']]
covid = covid.set_axis(['Date', 'StateCode', 'StateName', 'Pop', 'NewDeath',
'HospNow', 'NewCase', 'NewTest'], axis=1)

# Insert code that sets index with 'Date' column as a datetime object.
covid['Date'] = pd.to_datetime(covid['Date'], format="%Y-%m-%d")

covid.info()
covid
```

```python
grader.check("q5")
```

**Question 6 (3pt):** 

- Count the number of NaNs or 0s in the `NewCase` column of covid and store it in a variable named `num_no_case`. Hint: The `np.sum()` function or the 'Series object's `sum()` method  can be useful to count the number of `True`’s in a logical vector. Create a logical vector that represent whether its elements are `NaN` or 0 (i.e., `covid['NewCase'].isna() | covid['NewCase']==0`).

- Count the number of all `NaN`s in covid and store it in a variable named `num_NaN`. You may have to run `sum()` twice because `sum()` performs a summation of array elements over a given axis. 

- Run `covid.info()` and `covid.tail()` and see what columns have NAs. You may also notice that `Pop` is currently not a numeric type. Change the data type of `Pop` to numeric. Hint: `s.str.replace(",","")` can remove all commas in all elements in Series `s`. You may then want to use the `astype()` function after removing the commas.

- Replace all `NaN`s in column `HospNow` with 0s.

```python
# Write your code here ###
num_no_case = np.sum(covid['NewCase'].isna() | covid['NewCase']==0)
num_NaN = covid.isna().sum().sum()

covid['Pop']=covid['Pop'].str.replace(",","").astype(int)
covid.loc[covid['HospNow'].isna(),'HospNow']=0
covid.info()
covid.tail()
```

```python
grader.check("q6")
```

## Data Visualization with pandas and matplotlib

```python
# The following code is to show how to visualize your data in Python
import matplotlib.pyplot as plt
plt.close("all")

covid5 = covid.loc[covid['StateCode'].isin(['NC','VA','SC','TN','GA']), ['Date','StateCode','HospNow']]
covid5.set_index('Date',inplace=True)
covid5.groupby('StateCode')['HospNow'].plot(legend=True)
```

**Question 7 (1pt):** 
    
Modify `covid` to create a new column named `HospPerKiloPop`, which is calculated as `1000*HospNow/Pop`. This measure represents the number of hospitalized patients per 1,000 people, and thus it is a good measure to compare states with different sizes. Plot `HospPerKiloPop` instead of 'HospNow' to see whether hopitalization in North Carolina was better or worse than other neighboring states.

```python
# Write your code here ###
covid['HospPerKiloPop'] = 1000*covid['HospNow']/covid['Pop']
```

```python
covid5 = covid.loc[covid['StateCode'].isin(['NC','VA','SC','TN','GA']), ['Date','StateCode','HospPerKiloPop']]
covid5.set_index('Date',inplace=True)
covid5.groupby('StateCode')['HospPerKiloPop'].plot(legend=True)

del(covid['HospPerKiloPop']) # Delete the newly added column
```

```python
grader.check("q7")
```

**Question 8 (3pt):** 

Below you will find a code that generates a monthly summary data named `covid_mon`.  Add your code for the following:
- Add a new column named `PosRate` to `covid_mon`, which is calculated as `MCase/MTest`.
- Add a new column named `MDeathPerMilPop`, which represents monthly number of deaths per milllion people. 
- Add a set of new columns `MHospPerKiloPop`, `MCasePerKiloPop`, and `MTestPerKiloPop`. These should represent monthly total hopitalized days, number of cases, and number of tests per 1,000 people. 	

```python
#The following code generates a monthly summary data.
covid['Year'] = pd.DatetimeIndex(covid['Date']).year
covid['Month'] = pd.DatetimeIndex(covid['Date']).month
covid = covid.loc[((covid['Year']==2020) & (covid['Month']>=3))| ((covid['Year']==2021) & (covid['Month']<=2)),:]

covid_mon = covid.groupby(by=['Year','Month','StateCode'])[['NewDeath','HospNow','NewCase','NewTest']].sum()
covid_mon_add = covid.groupby(by=['Year','Month','StateCode'])[['Date','Pop']].max()
covid_mon = pd.concat([covid_mon, covid_mon_add],axis=1)
covid_mon = covid_mon.rename({'NewDeath':'MDeath','HospNow':'MHosp','NewCase':'MCase','NewTest':'MTest'}, axis=1)
covid_mon = covid_mon.reset_index()
```

```python
# Write your code here ###
covid_mon['PosRate'] = covid_mon['MCase']/covid_mon['MTest']
covid_mon['MDeathPerMilPop']=1000000*covid_mon.MDeath/covid_mon.Pop
covid_mon['MHospPerKiloPop']=1000*covid_mon.MHosp/covid_mon.Pop
covid_mon['MCasePerKiloPop']=1000*covid_mon.MCase/covid_mon.Pop
covid_mon['MTestPerKiloPop']=1000*covid_mon.MTest/covid_mon.Pop
```

```python
grader.check("q8")
```

### The codes below will let you visualize each COVID-19 metric.  
Feel free to mess with & improve the visualizations such as adding the title.  No points are assigned though.

```python
pos_rate = covid_mon.loc[covid_mon['StateCode'].isin(['NC','VA','SC','TN','GA']), ['Date','StateCode','PosRate']]
pos_rate.set_index('Date').groupby('StateCode')['PosRate'].plot(legend=True)
```

```python
mdeath = covid_mon.loc[covid_mon['StateCode'].isin(['NC','VA','SC','TN','GA']), ['Date','StateCode','MDeathPerMilPop']]
mdeath.set_index('Date').groupby('StateCode')['MDeathPerMilPop'].plot(legend=True)
```

```python
mhosp = covid_mon.loc[covid_mon['StateCode'].isin(['NC','VA','SC','TN','GA']), ['Date','StateCode','MHospPerKiloPop']]
mhosp.set_index('Date').groupby('StateCode')['MHospPerKiloPop'].plot(legend=True)
```

```python
mcase = covid_mon.loc[covid_mon['StateCode'].isin(['NC','VA','SC','TN','GA']), ['Date','StateCode','MCasePerKiloPop']]
mcase.set_index('Date').groupby('StateCode')['MCasePerKiloPop'].plot(legend=True)
```

```python
mtest = covid_mon.loc[covid_mon['StateCode'].isin(['NC','VA','SC','TN','GA']), ['Date','StateCode','MTestPerKiloPop']]
mtest.set_index('Date').groupby('StateCode')['MTestPerKiloPop'].plot(legend=True)
```

**Question 9 (2pt):** 
    
- Use `pivot_table()` so that StateCode values (e.g., ’NC’) become variables and each cell contains a `MHosp` value. You need `Date`, `StateCode`, and `MHosp` in `pivot_table()` to get the format you want. Assign the modified DataFrame to a variable named `hospital_days`.
- Modify `hospital_days` to add a new column named `USTotal` that aggregates the number of hospital days. Hint: Use the correct axis value.

```python
# Write your code here ###
hospital_days = covid_mon.pivot_table(values='MHosp', index='Date',columns="StateCode")
hospital_days['USTotal'] = hospital_days.sum(axis=1)
```

```python
# Check if your code works
hospital_days
```

```python
grader.check("q9")
```

**Honor Pledge:** As a signature for the following honor pledge, store your name as it is appeared on Canvas (e.g., "Mike Jordan") and Login ID (e.g., "mjorda21") in variables `my_name` and `my_id` below.  

***I affirm that I have not given or received any unauthorized help on this assignment and that this work is my own.***

```python
my_name = "Faye Struble"
my_id = "fstruble"
```

Tests for pledge (points: [0.5])


Tests for q1 (points: [1])
```
>>> loans['LoanStatus'].dtype=='category'
True
```


Tests for q2 (points: 4)
```
>>> ('LoanAmount' in loans.columns) and ('CurrentLoanAmount' not in loans.columns)
True
```
```
>>> 312313<avg_loan and avg_loan<312314
True
```
```
>>> np.mean(loans['LoanAmount'])>312313 and loans['LoanAmount'].mean()<312314
True
```


Tests for q3 (points: 3)
```
>>> 'Amt_per_Acct' in loans.columns.values
True
```
```
>>> loans.loc[2,'Amt_per_Acct']>17350 and loans.loc[2,'Amt_per_Acct']<17351
True
```
```
>>> all(fully_paid['LoanStatus']=='Fully Paid')
True
```


Tests for q4 (points: [0.5, 0.5, 0.5, 0.5])
```
>>> (high_risk.columns.values==['CustomerID', 'LoanAmount', 'Credit', 'Income', 'CreditScore']).all()
True
```
```
>>> high_risk.shape==(20754, 5)
True
```
```
>>> (valued_customer.columns.values==np.array(['LoanID', 'LoanStatus', 'OpenAccounts', 'Amt_per_Acct',
...        'MonthlyDebt'], dtype=object)).all()
True
```
```
>>> (valued_customer['LoanStatus']=='Fully Paid').all() and (valued_customer['Amt_per_Acct']>=125000).all()
True
```


Tests for q5 (points: 3)
```
>>> covid_all.shape==(20780, 45)
True
```
```
>>> covid_main.shape==(18995, 45)
True
```
```
>>> covid.columns.isin(['Date', 'StateCode', 'StateName', 'Pop', 'NewDeath', 'HospNow','NewCase', 'NewTest']).all()
True
```


Tests for q6 (points: [1, 1, 1])
```
>>> (not covid['HospNow'].isna().any())
True
```


Tests for q7 (points: [0.5, 0.5])
```
>>> (covid5['HospPerKiloPop'].sum()>315) and (covid5['HospPerKiloPop'].sum()<316)
True
```


Tests for q8 (points: 3)
```
>>> np.all(covid_mon.columns.values==np.array(['Year', 'Month', 'StateCode', 'MDeath', 'MHosp', 'MCase', 'MTest',
...        'Date', 'Pop', 'PosRate', 'MDeathPerMilPop', 'MHospPerKiloPop',
...        'MCasePerKiloPop', 'MTestPerKiloPop'], dtype=object))
True
```
```
>>> (covid_mon['PosRate'].mean()>0.090) and (covid_mon['PosRate'].mean()<0.091)
True
```
```
>>> (covid_mon['MDeathPerMilPop'].mean()>120) and (covid_mon['MDeathPerMilPop'].mean()<121)
True
```
```
>>> (covid_mon['MHospPerKiloPop'].mean()>4.5) and (covid_mon['MHospPerKiloPop'].mean()<4.6)
True
```
```
>>> (covid_mon['MCasePerKiloPop'].mean()>7.1) and (covid_mon['MCasePerKiloPop'].mean()<7.2)
True
```
```
>>> (covid_mon['MTestPerKiloPop'].mean()>93) and (covid_mon['MTestPerKiloPop'].mean()<94)
True
```


Tests for q9 (points: 2)
```
>>> (hospital_days.columns.values==np.array(['AK', 'AL', 'AR', 'AZ', 'CA', 'CO', 'CT', 'DC', 'DE', 'FL', 'GA',
...        'HI', 'IA', 'ID', 'IL', 'IN', 'KS', 'KY', 'LA', 'MA', 'MD', 'ME',
...        'MI', 'MN', 'MO', 'MS', 'MT', 'NC', 'ND', 'NE', 'NH', 'NJ', 'NM',
...        'NV', 'NY', 'OH', 'OK', 'OR', 'PA', 'RI', 'SC', 'SD', 'TN', 'TX',
...        'UT', 'VA', 'VT', 'WA', 'WI', 'WV', 'WY', 'USTotal'], dtype=object)).all()
True
```
```
>>> hospital_days['USTotal'].sum()==20223449
True
```
